Avant de soumettre cet exercice, assurez-vous que tout se déroule comme prévu. D'abord, **restart the kernel** (dans la barre de menu, sélectionnez Kernel$\rightarrow$Restart) et après **run all cells** (dans la barre de menu, sélectionnez Cell$\rightarrow$Run All).

Assurez-vous de remplir tous les endroit qui dit `YOUR CODE HERE` ou "YOUR ANSWER HERE", ainsi que votre nom ci-dessous:

```python
NAME = ""
```

---

## Dans cet exercice, nous allons nous entraîner avec NumPy. Pour chaque problème, implémentez la fonction demandée. Chacun compte pour un point sur vingt

```python
import numpy as np
```

<!-- grade_id: cell-0a37e1f21d2d54ee, points: 0, grade: false, solution: true -->
```python
def create_array(l: list) -> np.ndarray:
    """
    Function that converts a list into a NumPy array

    Parameters
    ----------
    l : list
        A list containing a given number of elements

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-550cdda2e6fb262c, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(create_array([1, 2, 3, 4, 5, 6, 7, 8, 9, 10]), np.arange(1, 11, 1))
a = np.random.rand(10,1)
assert np.array_equal(create_array(a.data), a)
```

<!-- grade_id: cell-75c2b8c4eab92052, points: 0, grade: false, solution: true -->
```python
def zeros(n: int) -> np.ndarray:
    """
    Create a numpy array of all zeros of size n

    Parameters
    ----------
    n : int
        Size of the created array

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-86db41a3dd1aacc6, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(zeros(3), [0, 0, 0])
assert np.array_equal(zeros(0), [])
```

<!-- grade_id: cell-f5163451497f11f4, points: 0, grade: false, solution: true -->
```python
def bigger_than(a: np.ndarray, n: int) -> np.ndarray:
    """
    Create a view of the elements bigger than n

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)
    n : int
        The value used for comparison

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-14936f0983dd2892, points: 1, grade: true, solution: false -->
```python
arr = np.array([1, 2, 3, 4, 5, 6])
assert np.array_equal(bigger_than(arr, 3), [4, 5, 6])
arr = np.loadtxt("data/testa.txt")
assert np.array_equal(bigger_than(arr, 3), [7, 7, 8, 5, 7, 8])
```

<!-- grade_id: cell-4c5932021ce6e2c0, points: 0, grade: false, solution: true -->
```python
def even(a: np.ndarray) -> np.ndarray:
    """
    Create a view of the elements of the array a that are even

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2206870c64ab8070, points: 1, grade: true, solution: false -->
```python
arr = np.array([1, 2, 3, 4, 5, 6])
assert np.array_equal(even(arr), [1, 3, 5])
arr = np.loadtxt("data/testa.txt")
assert np.array_equal(even(arr), [7, 7, 1.0, 5.0, 3.0, 3.0, 7.0, 3.0])
```

<!-- grade_id: cell-1a730de84ad54475, points: 0, grade: false, solution: true -->
```python
def replace_bigger_than(a: np.ndarray, n: int) -> np.ndarray:
    """
    Replace values that are bigger than n by -1

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)
    n : int
        The number used for comparison

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-cccb4841ace09ada, points: 1, grade: true, solution: false -->
```python
arr = np.array([1, 2, 3, 4, 5, 6])
assert np.array_equal(replace_bigger_than(arr, 3), [1, 2, 3, -1, -1, -1])
arr = np.loadtxt("data/testa.txt")
assert np.array_equal(replace_bigger_than(arr, 3), [-1, -1, -1, 1.0, -1, 3, 3, -1, 3, -1])
```

<!-- grade_id: cell-e53d12706a5b8e81, points: 0, grade: false, solution: true -->
```python
def replace_bigger_than_copy(a: np.ndarray, n: int) -> np.ndarray:
    """
    Replace values that are bigger than n by -1 without modifying the original array a

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)
    n : int
        The number used for comparison

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6c036e0a7b445bc3, points: 1, grade: true, solution: false -->
```python
arr = np.array([1, 2, 3, 4, 5, 6])
assert np.array_equal(replace_bigger_than_copy(arr, 3), [1, 2, 3, -1, -1, -1])
assert np.array_equal(arr, [1, 2, 3, 4, 5, 6])
arr = np.loadtxt("data/testa.txt")
assert np.array_equal(replace_bigger_than_copy(arr, 3), [-1, -1, -1, 1.0, -1, 3, 3, -1, 3, -1])
assert np.array_equal(arr, np.loadtxt("data/testa.txt"))
```

<!-- grade_id: cell-63df1914534e3682, points: 0, grade: false, solution: true -->
```python
def repeat_array(a: np.ndarray, n: int) -> np.ndarray:
    """
    Create an array that is the repetition of array a n times

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)
    n : int
        How many times to repeat the array

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0e884597194730d1, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(repeat_array(np.array([1, 2, 3]), 3), [1, 2, 3, 1, 2, 3, 1, 2, 3])
assert np.array_equal(repeat_array(np.loadtxt("data/testa.txt"), 3), [7.0, 7.0, 8.0, 1.0, 5.0, 3.0, 3.0, 7.0, 3.0, 8.0, 7.0, 7.0, 8.0, 1.0, 5.0, 3.0, 3.0, 7.0, 3.0, 8.0, 7.0, 7.0, 8.0, 1.0, 5.0, 3.0, 3.0, 7.0, 3.0, 8.0])
```

<!-- grade_id: cell-b18b08bda10520b6, points: 0, grade: false, solution: true -->
```python
def common_items(a: np.ndarray, b: np.ndarray) -> np.ndarray:
    """
    Create an array that contains common items in arrays a and b

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)
    b : ndarray
        An array of numbers (float / int)

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e651547944e27e4a, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(common_items(np.array([1, 2, 3, 4, 5]), np.array([6, 7, 8, 9, 10])), [])
assert np.array_equal(common_items(np.loadtxt("data/testa.txt"), np.loadtxt("data/testb.txt")), [3, 8])
```

<!-- grade_id: cell-f8cf1f96e4f93ef3, points: 0, grade: false, solution: true -->
```python
def where_equal(a: np.ndarray, b: np.ndarray) -> np.ndarray:
    """
    Create an array with the positions of the elements in a and b that are equal 

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)
    b : ndarray
        An array of numbers (float / int)

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f189c9173e6baf32, points: 1, grade: true, solution: false -->
```python
assert where_equal(np.array([1, 2, 3, 4, 5, 6, 7, 8, 9]), np.array([9, 8, 7, 6, 5, 4, 3, 2 ,1]))[0] == 4
assert np.array_equal(where_equal(np.loadtxt("data/testa.txt"), np.loadtxt("data/testb.txt"))[0], [6, 9])
```

<!-- grade_id: cell-6ad80f6f187f4290, points: 0, grade: false, solution: true -->
```python
def count_negative_values(a: np.ndarray) -> int:
    """
    Calculate how many negative values there are in an array

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)

    Returns
    -------
    int
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-44650e1c8431ba74, points: 1, grade: true, solution: false -->
```python
assert count_negative_values(np.array([1, -1, 2, 3, 5, -2, -10.6])) == 3
assert count_negative_values(np.array([1, 1, 2, 3, 5, 2, 10.6])) == 0
```

<!-- grade_id: cell-2c3954d4eb9b42dd, points: 0, grade: false, solution: true -->
```python
def normalize(a: np.ndarray) -> ndarray:
    """
    Function that normalizes an array, i.e., produced values that are
    contained within 0 and 1.

    Parameters
    ----------
    a : ndarray
        An array of numbers (float / int)

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-390e123d17f5b230, points: 3, grade: true, solution: false -->
```python
data = np.loadtxt("data/seplength.txt")
assert np.array_equal(normalize(data), [0.22222222222222213, 0.1666666666666668, 0.11111111111111119, 0.08333333333333327, 0.19444444444444448, 0.30555555555555564, 0.08333333333333327, 0.19444444444444448, 0.027777777777777922, 0.1666666666666668, 0.30555555555555564, 0.13888888888888887, 0.13888888888888887, 0.0, 0.41666666666666663, 0.38888888888888895, 0.30555555555555564, 0.22222222222222213, 0.38888888888888895, 0.22222222222222213, 0.30555555555555564, 0.22222222222222213, 0.08333333333333327, 0.22222222222222213, 0.13888888888888887, 0.19444444444444448, 0.19444444444444448, 0.25000000000000006, 0.25000000000000006, 0.11111111111111119, 0.13888888888888887, 0.30555555555555564, 0.25000000000000006, 0.3333333333333333, 0.1666666666666668, 0.19444444444444448, 0.3333333333333333, 0.1666666666666668, 0.027777777777777922, 0.22222222222222213, 0.19444444444444448, 0.055555555555555594, 0.027777777777777922, 0.19444444444444448, 0.22222222222222213, 0.13888888888888887, 0.22222222222222213, 0.08333333333333327, 0.27777777777777773, 0.19444444444444448, 0.7499999999999999, 0.5833333333333334, 0.7222222222222222, 0.3333333333333333, 0.611111111111111, 0.38888888888888895, 0.5555555555555555, 0.1666666666666668, 0.6388888888888887, 0.25000000000000006, 0.19444444444444448, 0.44444444444444453, 0.4722222222222222, 0.4999999999999999, 0.361111111111111, 0.6666666666666666, 0.361111111111111, 0.41666666666666663, 0.5277777777777778, 0.361111111111111, 0.44444444444444453, 0.4999999999999999, 0.5555555555555555, 0.4999999999999999, 0.5833333333333334, 0.6388888888888887, 0.6944444444444443, 0.6666666666666666, 0.4722222222222222, 0.38888888888888895, 0.3333333333333333, 0.3333333333333333, 0.41666666666666663, 0.4722222222222222, 0.30555555555555564, 0.4722222222222222, 0.6666666666666666, 0.5555555555555555, 0.361111111111111, 0.3333333333333333, 0.3333333333333333, 0.4999999999999999, 0.41666666666666663, 0.19444444444444448, 0.361111111111111, 0.38888888888888895, 0.38888888888888895, 0.5277777777777778, 0.22222222222222213, 0.38888888888888895, 0.5555555555555555, 0.41666666666666663, 0.7777777777777776, 0.5555555555555555, 0.611111111111111, 0.9166666666666665, 0.1666666666666668, 0.8333333333333333, 0.6666666666666666, 0.8055555555555556, 0.611111111111111, 0.5833333333333334, 0.6944444444444443, 0.38888888888888895, 0.41666666666666663, 0.5833333333333334, 0.611111111111111, 0.9444444444444444, 0.9444444444444444, 0.4722222222222222, 0.7222222222222222, 0.361111111111111, 0.9444444444444444, 0.5555555555555555, 0.6666666666666666, 0.8055555555555556, 0.5277777777777778, 0.4999999999999999, 0.5833333333333334, 0.8055555555555556, 0.8611111111111112, 1.0, 0.5833333333333334, 0.5555555555555555, 0.4999999999999999, 0.9444444444444444, 0.5555555555555555, 0.5833333333333334, 0.4722222222222222, 0.7222222222222222, 0.6666666666666666, 0.7222222222222222, 0.41666666666666663, 0.6944444444444443, 0.6666666666666666, 0.6666666666666666, 0.5555555555555555, 0.611111111111111, 0.5277777777777778, 0.44444444444444453])
```

<!-- grade_id: cell-2, points: 0, grade: false, solution: true -->
```python
def text_arrays_to_unified(a: np.ndarray) -> int:
    """
    Function that converts the array contained in the 'data/arrays.txt' 
    file and returns a single array (i.e., shape (x, 1)) of all 
    separate values 

    Parameters
    ----------
    a : ndarray
        An array of strings

    Returns
    -------
    ndarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2test, points: 3, grade: true, solution: false -->
```python
data = np.loadtxt("data/arrays.txt", dtype=str)

assert np.array_equal(text_arrays_to_unified(data), [3, 2, 1, 1, 1, 2, 0, 4, 1, 0, 0, 4, 0, 3, 1, 4, 2, 0, 1, 0, 3, 3, 2, 2, 0, 0, 0, 0, 3, 3, 3, 1, 4, 0, 3, 4, 1, 3, 1, 2])
```

<!-- grade_id: cell-3dsadas, points: 3, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt

def plot_histogram(x: np.ndarray):
    """
    Plot an histogram of the data contained in the array x. 
    Do not forget to properly annotate the plot (axis, title, etc.)

    Parameters
    ----------
    x : ndarray
        An array of numbers (float / int)

    Returns
    -------
    None
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_histogram(np.loadtxt("data/x.txt"))
```

### By visual inspection, what type of data distribution follows the data you plotted in the previous exercise? 

Answer here: